# Question 2

In this question, we'll dive a bit deeper into some arithmetic operations.

### Part A

Five numbers are given below. Keeping them in that same order, compute a single result with Python addition, multiplication, subtraction, and division, respectively.

Put another way--add the first and second numbers, multiply the second and third numbers, subtract the third and fourth numbers, and divide the fourth and fifth numbers. Don't use any parentheses; let order of operations take care of things. Do this in a single statement, and store the result in the `result` variable below.

<!-- grade_id: q2a, points: 10, grade: false, solution: true -->
```python
n1 = 0.4589
n2 = 13
n3 = -59
n4 = 8.43333333
n5 = 2.1

result = -1

### BEGIN SOLUTION
result=n1+n2*n3-n4/n5
result #checking solution
### END SOLUTION
```

Output:
```
-770.5569730142857
```

<!-- grade_id: q2a_test, points: 7, grade: true, solution: false -->
```python
import numpy as np
np.testing.assert_allclose(result, -770.5569730142857)
```

Output:
```
NameError                                 Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_11572/588583021.py in <module>
      1 import numpy as np
----> 2 np.testing.assert_allclose(result, -770.5569730142857)

NameError: name 'result' is not defined
```

### Part B

Using the same five numbers as before, this time use parentheses around your arithmetic operation to enforce that the *addition* and *subtraction* operations should happen before the others. Again, store your result in the variable `result`.

<!-- grade_id: q2b, points: 0, grade: false, solution: true -->
```python
n1 = 0.4589
n2 = 13
n3 = -59
n4 = 8.43333333
n5 = 2.1

result = -1

### BEGIN SOLUTION
result=(n1+n2)*(n3-n4)/n5
result  #checking solution
### END SOLUTION
```

Output:
```
-432.1802333119699
```

<!-- grade_id: q2b_test, points: 7, grade: true, solution: false -->
```python
import numpy as np
np.testing.assert_allclose(result, -432.1802333119699)
```

### Part C

Repeat the steps you took in Part B, this time casting each variable to an integer data type before performing any computation. Store your answer in the variable `result`.

<!-- grade_id: q2c, points: 0, grade: false, solution: true -->
```python
n1 = 0.4589
n2 = 13
n3 = -59
n4 = 8.43333333
n5 = 2.1

result = -1

### BEGIN SOLUTION
n1=int(n1)
n2=int(n2)
n3=int(n3)
n4=int(n4)
n5=int(n5)     #casting variables as integers
result=(n1+n2)*(n3-n4)/n5     #arithmetic operation
result   #checking solution
### END SOLUTION
```

Output:
```
-435.5
```

<!-- grade_id: q2c_test, points: 7, grade: true, solution: false -->
```python
assert result == -435.5
```

### Part D

Two numbers are given below. When you perform the operation `n1 / n2`, what is the quotient? Store this quantity in the variable named `quotient`.

<!-- grade_id: q2d, points: 0, grade: false, solution: true -->
```python
n1 = 20344983945038.48374839
n2 = 984374.783894273

quotient = -1

### BEGIN SOLUTION
quotient=n1/n2
quotient  #checking solution
### END SOLUTION
```

Output:
```
20667924.735487375
```

<!-- grade_id: q2d_test, points: 7, grade: true, solution: false -->
```python
import numpy as np
np.testing.assert_allclose(quotient, 20667924.735487375)
```

### Part E

You've seen in lecture how, when two numbers in Python are divided, their result is a float. What if you *really* wanted the output of division to be an integer? See if you can make that happen by modifying the following single line of code.

(You are **not** allowed to add any numbers or additional lines of code!)

<!-- grade_id: q2e, points: 0, grade: false, solution: true -->
```python
integer_quotient = 32 / 16

### BEGIN SOLUTION
integer_quotient=int(integer_quotient)  #casting
integer_quotient  #checking solution
### END SOLUTION
```

Output:
```
2
```

<!-- grade_id: q2e_test, points: 7, grade: true, solution: false -->
```python
assert type(integer_quotient) == int
```